# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit. 

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point. 

There are 5 total possible points in this sprint challenge. 


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section). 
____



### Import Data

```python
# # could use *web_lg or *web_sm instead
# !python -m spacy download en_core_web_md
```

Output:
```
Collecting en_core_web_md==2.2.5
  Downloading https://github.com/explosion/spacy-models/releases/download/en_core_web_md-2.2.5/en_core_web_md-2.2.5.tar.gz (96.4 MB)
[K     |████████████████████████████████| 96.4 MB 1.2 MB/s 
[?25hRequirement already satisfied: spacy>=2.2.2 in /usr/local/lib/python3.7/dist-packages (from en_core_web_md==2.2.5) (2.2.4)
Requirement already satisfied: wasabi<1.1.0,>=0.4.0 in /usr/local/lib/python3.7/dist-packages (from spacy>=2.2.2->en_core_web_md==2.2.5) (0.8.2)
Requirement already satisfied: blis<0.5.0,>=0.4.0 in /usr/local/lib/python3.7/dist-packages (from spacy>=2.2.2->en_core_web_md==2.2.5) (0.4.1)
Requirement already satisfied: plac<1.2.0,>=0.9.6 in /usr/local/lib/python3.7/dist-packages (from spacy>=2.2.2->en_core_web_md==2.2.5) (1.1.3)
Requirement already satisfied: numpy>=1.15.0 in /usr/local/lib/python3.7/dist-packages (from spacy>=2.2.2->en_core_web_md==2.2.5) (1.19.5)
Requirement already satisfied: requests<3.0.0,>=2.13.0 in /usr/local/lib/p
```

```python
# !wget https://raw.githubusercontent.com/LambdaSchool/DS-Unit-4-Sprint-1-NLP/main/requirements.txt
# !pip install -r requirements.txt
```

Output:
```
--2021-10-01 19:27:06--  https://raw.githubusercontent.com/LambdaSchool/DS-Unit-4-Sprint-1-NLP/main/requirements.txt
Resolving raw.githubusercontent.com (raw.githubusercontent.com)... 185.199.111.133, 185.199.109.133, 185.199.108.133, ...
Connecting to raw.githubusercontent.com (raw.githubusercontent.com)|185.199.111.133|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 149 [text/plain]
Saving to: ‘requirements.txt’

requirements.txt    100%[===================>]     149  --.-KB/s    in 0s      

2021-10-01 19:27:06 (3.32 MB/s) - ‘requirements.txt’ saved [149/149]

Collecting gensim==3.8.1
  Downloading gensim-3.8.1-cp37-cp37m-manylinux1_x86_64.whl (24.2 MB)
[K     |████████████████████████████████| 24.2 MB 74 kB/s 
[?25hCollecting pyLDAvis==2.1.2
  Downloading pyLDAvis-2.1.2.tar.gz (1.6 MB)
[K     |████████████████████████████████| 1.6 MB 50.4 MB/s 
[?25hCollecting spacy==2.2.3
  Downloading spacy-2.2.3-cp37-cp37m-manylinux1_x86_64.whl (10.4 MB)
[K     |██
```

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
# YOUR CODE HERE

df = pd.read_json(data_url, lines = True)
```

```python
df
```

Output:
```
business_id  cool  ... useful                 user_id
0     nDuEqIyRc8YKS1q1fX0CZg     1  ...     10  n1LM36qNg4rqGXIcvVXv8w
1     eMYeEapscbKNqUDCx705hg     0  ...      0  5CgjjDAic2-FAvCtiHpytA
2     6Q7-wkCPc1KF75jZLOTcMw     1  ...      2  BdV-cf3LScmb8kZ7iiBcMA
3     k3zrItO4l9hwfLRwHBDc9w     3  ...      5  cZZnBqh4gAEy4CdNvJailQ
4     6hpfRwGlOzbNv7k5eP9rsQ     1  ...      5  n9QO4ClYAS7h9fpQwa5bhA
...                      ...   ...  ...    ...                     ...
9995  1h3ysSuSazvXc1aeLiiOew     0  ...      1  QFYqAk8n5Z1O3t7zwjA7Hg
9996  Rwahe1zbFpw6VIjb5ngZeg     0  ...      0  X7jQ-4788irfe5ABZNvYcA
9997  8itGZAOBMiTbHKOwLuh4_Q     0  ...      0  _pi5J_1CIQWceLhTJkx_yA
9998  A5Rkh7UymKm0_Rxm9K2PJw     0  ...      0  PP1K311ZKbpDgTjwic3u5Q
9999  xtYiHTmunjfCN2sUaQxBjA     0  ...      0  Y1tHSar3k-_clZVCM7JBZQ

[10000 rows x 9 columns]
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.
import spacy
nlp = spacy.load('en_core_web_sm')

def tokenize(doc):
# YOUR CODE HERE

    """
    Takes text and returns a list of tokens in the form of lemmas.
    This will not include puncuation, or stop words.
    """
    
    doc = nlp(doc)
    
    return [token.lemma_.strip() for token in doc if (token.is_stop != True) and (token.is_punct != True)]

df["lemmas"] = df['text'].apply(tokenize)
```

```python
df['lemmas']
```

Output:
```
0       [beware, fake, FAKE, fake, small, business, Lo...
1       [come, lunch, Togo, service, quick, staff, fri...
2       [Vegas, dozen, time, step, foot, Circus, Circu...
3       [go, night, close, street, party, good, actual...
4       [3.5, 4, star, , bad, price, $, 12.99, lunch, ...
                              ...                        
9995    [family, hungry, Subway, open, 24, hour, guy, ...
9996    [wife, come, couple, friend, sever, excited, p...
9997    [food, ok, brag, food, hot, item, tasty, horri...
9998    [today, visit, great, love, enjoy, Town, Squar...
9999    [absolute, bad, place, stay, 43, year, life, ,...
Name: lemmas, Length: 10000, dtype: object
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this. 

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# Create a vector representation of the reviews 
# Name that doc-term matrix "dtm"

# YOUR CODE HERE

from sklearn.feature_extraction.text import TfidfVectorizer

tfidf = TfidfVectorizer(ngram_range=(1,2),
                        max_df=.97,
                        min_df=.02,
                        tokenizer=tokenize)
                        
dtm = tfidf.fit_transform(df['lemmas'][0])
dtm = pd.DataFrame(dtm.todense(), columns=tfidf.get_feature_names())
```

Output:
```
/usr/local/lib/python3.7/dist-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/usr/local/lib/python3.7/dist-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/usr/local/lib/python3.7/dist-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/usr/local/lib/python3.7/dist-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/usr/local/lib/python3.7/dist-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use sele
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

# YOUR CODE HERE
nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree')
nn.fit(dtm)
```

Output:
```
NearestNeighbors(algorithm='kd_tree', leaf_size=30, metric='minkowski',
                 metric_params=None, n_jobs=None, n_neighbors=10, p=2,
                 radius=1.0)
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

# YOUR CODE HERE

# fake review
sample = 'This product is awesome!'
sample_dtm = tfidf.transform(pd.Series([sample]))
sample_dtm = pd.DataFrame(sample_dtm.todense(), columns=tfidf.get_feature_names())

# Query Using kneighbors
neigh_dist, neigh_index = nn.kneighbors(sample_dtm)
for x in neigh_index:
  print(df.iloc[x]['text'])
```

Output:
```
4     3.5 to 4 stars\n\nNot bad for the price, $12.9...
13    The food was great!  It was super busy but our...
12    Great friendly customer service and quality fo...
8     Absolutely the most Unique experience in a nai...
10    We popped in for dinner yesterday with no rese...
11    Thw worst stay ever! So first i ended up payin...
6     This show is absolutely amazing!! What an incr...
0     BEWARE!!! FAKE, FAKE, FAKE....We also own a sm...
9     Wow. I walked in and sat at the bar for 10 min...
7     Came for the Pho and really enjoyed it!  We go...
Name: text, dtype: object
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset. 

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels). 
    - Use that Pipeline to predict a star rating for your fake review from Part 2. 



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`. 
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1** 
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import GridSearchCV

# Name the gridsearch instance "gs"

# YOUR CODE HERE

target = 'stars'

y, X = df[target], df.drop(columns=target)

X = [list(v)[0] for v in X.values]

tfidf = TfidfVectorizer(stop_words="english", tokenizer=None)

knc = KNeighborsClassifier()

pipe = Pipeline([("vect", tfidf), 
                 ("clf", knc)])  
parameters = {
    'vect__max_df': (0.75, 1.0),
    'vect__max_features': (500, 1000),
    'clf__algorithm':(['ball_tree','brute']),
    'clf__weights':(['uniform','distance']),
}
gs = GridSearchCV(pipe,parameters, n_jobs=1)
gs.fit(X,y)
```

Output:
```
GridSearchCV(cv=None, error_score=nan,
             estimator=Pipeline(memory=None,
                                steps=[('vect',
                                        TfidfVectorizer(analyzer='word',
                                                        binary=False,
                                                        decode_error='strict',
                                                        dtype=<class 'numpy.float64'>,
                                                        encoding='utf-8',
                                                        input='content',
                                                        lowercase=True,
                                                        max_df=1.0,
                                                        max_features=None,
                                                        min_df=1,
                                                        ngram_range=(1, 1),
                                                       
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract. 
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this: 

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.** 

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz). 

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
from gensim.models.ldamodel import LdaModel
import gensim
import re
# import pyLDAvis
```

Output:
```
INFO: Pandarallel will run on 10 workers.
INFO: Pandarallel will use Memory file system to transfer data between the main process and workers.
/usr/local/lib/python3.7/dist-packages/past/types/oldstr.py:5: DeprecationWarning: Using or importing the ABCs from 'collections' instead of from 'collections.abc' is deprecated since Python 3.3,and in 3.9 it will stop working
  from collections import Iterable
```

```python
# from pandarallel import pandarallel
# # we must initalize pandarallel before we can use it
# pandarallel.initialize(progress_bar=True, nb_workers=10)
# # so that the progress bars will work
# from pandarallel.utils import progress_bars
# progress_bars.is_notebook_lab = lambda : True
```

### 1. Estimate a LDA topic model of the review tex

```python
def clean_text(text):
    """
    Accepts a single text document and performs several regex substitutions in order to clean the document. 
    
    Parameters
    ----------
    text: string or object 
    
    Returns
    -------
    text: string or object
    """
    
    # order of operations - apply the expression from top to bottom
    email_regex = r"From: \S*@\S*\s?"
    non_alpha = '[^a-zA-Z]'
    multi_white_spaces = "[ ]{2,}"
    
    text = re.sub(email_regex, "", text)
    text = re.sub(non_alpha, ' ', text)
    text = re.sub(multi_white_spaces, " ", text)
    
    # apply case normalization 
    return text.lower().lstrip().rstrip()
```

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value 
num_topics = 5

# use tokenize function you created earlier to create tokens 

# create a id2word object (hint: use corpora.Dictionary)

# create a corpus object (hint: id2word.doc2bow)

# instantiate an lda model

# YOUR CODE HERE

def filter_lemmas(lemmas):
    """
    Filter out any lemmas that are 2 characters or smaller
    """
    return [lemma for lemma in lemmas if len(lemma) > 2]

df["filtered_lemmas"] = df["lemmas"].parallel_map(filter_lemmas)

df['lemmas'] = df["filtered_lemmas"]

id2word = corpora.Dictionary(df['lemmas'])

corpus = [id2word.doc2bow(text) for text in df['lemmas']]

lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )
```

Output:
```
VBox(children=(HBox(children=(IntProgress(value=0, description='0.00%', max=1000), Label(value='0 / 1000'))), …
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
import seaborn as sns
import matplotlib.pyplot as plt

# # Use pyLDAvis (or a ploting tool of your choice) to visualize your results 

# # YOUR CODE HERE

# from gensim import corpora
import pyLDAvis.gensim
# # Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
# from gensim.models.ldamodel import LdaModel
# import gensim

pyLDAvis.enable_notebook()
vis = pyLDAvis.gensim.prepare(lda, corpus, id2word)
vis
```

Output:
```
PreparedData(topic_coordinates=              x         y  topics  cluster       Freq
topic                                                
4      0.035503 -0.065128       1        1  30.559925
0     -0.075416 -0.003740       2        1  29.655779
2     -0.063993  0.041863       3        1  19.785688
1      0.022265 -0.042177       4        1  14.112560
3      0.081642  0.069182       5        1   5.886048, topic_info=      Term         Freq        Total Category  logprob  loglift
137   food  4928.000000  4928.000000  Default  30.0000  30.0000
100   room  1072.000000  1072.000000  Default  29.0000  29.0000
354  order  3482.000000  3482.000000  Default  28.0000  28.0000
360  pizza   744.000000   744.000000  Default  27.0000  27.0000
197  great  4152.000000  4152.000000  Default  26.0000  26.0000
..     ...          ...          ...      ...      ...      ...
43    good   106.166374  6239.627328   Topic5  -5.5709  -1.2411
586  drink    73.947792  1378.896320   Topic5  -5.9325  -0.0931
37 
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

Groups 3 and 2 appear to be somewhat related. Their context appears to include quite a bit regarding restaurants. Grouping 3 looks to be more centered around deserts, ice-cream, and cheese pizza. While Grouping 2 seems to be centered about the review of the restaurant with postive remarks about the establishment.

Groups 4 and 1 appear to be more closely related. Their context appears to include a lot about buisnesses. Grouping 4 looks to be centered around a doctor's office, while grouping 1 looks to be centered around service, appointments, and customers. 

Group 5 appears to be related to home decorations, or furniture.